---
title       : Insert the chapter title here
description : Insert the chapter description here
---

## <<<New Exercise>>>

```yaml
type: NormalExercise
key: 4c919a781d
lang: r
xp: 100
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
