---
title       : Manipuler du texte en langage naturel
description : Insert the chapter description here


---
## Tokenisation

```yaml
type: NormalExercise
key: 4b7358efd0
lang: r
xp: 100
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
tib_ingredients=read.table("tib_ingredients.csv",row.names=FALSE)
```

`@solution`
```{r}
tib_ingredients=read.table("tib_ingredients.csv",row.names=FALSE)

```

`@sct`
```{r}
test_error()
```

---
## Retirer les mots-outils

```yaml
type: NormalExercise
key: dd7c384372
lang: r
xp: 100
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```


