---
title       : Visualisation
description : Insert the chapter description here


---
## Ingredients: occurrence

```yaml
type: NormalExercise
key: 615077e0c5
lang: r
xp: 100
skills: 1
```


`@instructions`

On va maintenant 

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
tib_ingredients=read.table("tib_ingredients.rds", row.names=FALSE)
```

`@solution`
```{r}
tib_ingredients=read.table("tib_ingredients.rds", row.names=FALSE)
```

`@sct`
```{r}
test_error()
```
